---
title       : Insert the chapter title here
description : Insert the chapter description here
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf



--- type:NormalExercise lang:python xp:100 skills:2 key:eac6c2e6a5
## Normal Exercise


*** =instructions

*** =hint

*** =pre_exercise_code
```{python}
pwd
```

*** =sample_code
```{python}

```

*** =solution
```{python}
pwd
```

*** =sct
```{python}

```
